This is a very simple example of an assignment notebook. 

Let us imagine that we want all students in a course to receive their **individual dataset** to work with. We use the `student_id` as a **unique identifier** of the student and their dataset. In our case the **student_id** is `123`, therefor their personal **dataset** is called `123.csv`. It is stored in a **folder** called `data`.


```python
student_id = '123'
```

## Task

Implement a function `import_dataset`, which imports the dataset assigned and provided to you as a CSV file using `pandas`. The function `import_dataset` must take a single parameter that holds the basename of the CSV file (i.e., the name without file extension). The output of the function `import_dataset` should be a dataframe `data`, which represents your given CSV file.


-----
*Explaination:*

*The task of the student is to correctly load the dataset into a dataframe. In the autograded test, we would like to compare the shape of the dataframe, created by the student, to the one we expect. 
Due to the fact that each student has an individual dataset, **it may occur that the shape of the datasets varies slightly** (e.g. different number of rows). This is why we want to run the **assert_equal** function,* 
* *using the **student solution** on the one hand and*
* *calling an **"answers module"** on the other hand. The "answers module" consists of functions that we want to validate the students answers against.*

*This way we make sure that the student answers are compared to an expected answer, based on the personal dataset they were given.* 

<!-- grade_id: cell-47266c307c418d73, points: 0, grade: false, solution: true -->
```python
def import_dataset(x):
### BEGIN SOLUTION
    import pandas as pd
    path_to_file = "data/" + str(x) + ".csv"
    data = pd.read_csv(path_to_file)
    
### END SOLUTION
    return data
```

<!-- grade_id: cell-767c64ac20ed870c, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
import pandas
# make all tests that are the same for all students visible (e.g. that they have created a pandas dataframe)
assert_equal(type(import_dataset(student_id)), pandas.core.frame.DataFrame)

### BEGIN HIDDEN TESTS
import importlib.util
import os
# call the created "answers module" using an absolute path (consider that the location and structure of the autograding function of nbgrader is different) in order to run the personalised evaluations
spec = importlib.util.spec_from_file_location("hw1", os.path.expanduser("~/eval_files/homework1.py"))
hw1 = importlib.util.module_from_spec(spec)
spec.loader.exec_module(hw1)
# run the assert_equal test by calling the module and the according function
assert_equal(import_dataset(student_id).shape, hw1.import_dataset(student_id).shape)
### END HIDDEN TESTS
```